### Before you start...

Before you start solving the challenge, press `Ctrl-Enter` (PC) or `Command-Enter` (Mac) in the cell below. It will load the modules needed for testing and auto-grading the challenges. After you did, a number between the square brackets will appear in front of the cell, like - for example:
```python
[1]
```
If successful, the following message will appear below the first cell:
```python
Assertions imported OK
Test cells imported OK
```

```python
# Run this cell first!

import httpimport
with httpimport.github_repo('eur-nl', 'assertions'):
    from assertions import *
with httpimport.github_repo('jjengelberts', 'precourse-test'):
    from test_cells import *
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[1], line 3
      1 # Run this cell first!
----> 3 import httpimport
      4 with httpimport.github_repo('eur-nl', 'assertions'):
      5     from assertions import *

ModuleNotFoundError: No module named 'httpimport'
```

<br>
<br>
<br>
<br>

## Challenge 2a: Let's Play Chess!
One of the stronger pieces in chess is the rook. Rooks move along rows or columns. In Figure 1, you will see a white rook on D5 with crosses horizontally and vertically indicating the possible squares it can move to. 
<div>
    <img src="https://raw.githubusercontent.com/jjengelberts/precourse-test/main/figures/rookmoves.svg" align="left" width="450"/><br clear="all">
    <b>Figure 1</b>
</div>
<br>
You will complete function `challenge_2a(square)` that outputs all possible moves of the rook given the current square. 

The function takes as argument the `square` the rook is presently on, and returns a list named `possible` containing all possible squares the rook can move to (on an otherwise empty board). The list is returned in **sorted** order, you don't have to sort it yourself.

**TIP**: If you loop through rows and columns to get the `possible` list, you add the current square the rook is on. Be aware of this and make sure **not** to include the given square in your `possible` moves list! 

If you are able, try to apply a [pythonic way](https://realpython.com/lessons/pythonic-loops-overview/) to solve this.

## Example input
```python
D5
```

## Example return value
```python
['A5', 'B5', 'C5', 'D1', 'D2', 'D3', 'D4', 'D6', 'D7', 'D8', 'E5', 'F5', 'G5', 'H5']
```

<!-- grade_id: code-chal-02a, points: 0, grade: false, solution: true -->
```python
# Code cell

def challenge_2a(square):

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Test cell

test_challenge_2a(challenge_2a)
```

<!-- grade_id: assert-chal-02a, points: 10, grade: true, solution: false -->
```python
# Assertion cell

check_assertion(challenge_2a)
```

## Next challenge

Great job! If you solved this challenge, you are ready for **Challenge 3a**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge3a.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 3a"/></a>

If you were not able to complete the challenge above, you can practice your skills a bit more by **Challenge 3b**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge3b.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 3b"/></a>