# Off We Go!

You are ready to start the first challenge. The first cell loads the modules needed for testing and auto-grading the challenges. To run the cell, press `Ctrl-Enter` (PC) or `Command-Enter` (Mac) in the cell below. 

If everything is corrrect, a number between the square brackets will appear in front of the cell, like - for example:
```python
[1]
```
If the modules are successfully loaded, the following message will appear below the first cell:
```python
Assertions imported OK
Test cells imported OK
```

```python
# Run this cell to load the modules!

import httpimport
with httpimport.github_repo('eur-nl', 'assertions'):
    from assertions import *
with httpimport.github_repo('jjengelberts', 'precourse-test'):
    from test_cells import *
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[2], line 3
      1 # Run this cell first!
----> 3 import httpimport
      4 with httpimport.github_repo('eur-nl', 'assertions'):
      5     from assertions import *

ModuleNotFoundError: No module named 'httpimport'
```

<br> ### IETS OVER ALS HET NIET LUKT? ###
<br>
<br>

## Challenge 1

In this first challenge, you will complete a function that finds the highest number. 

The function is called `challenge_1(a, b, c)` and should compare the three arguments `a`, `b` and `c` (these should all be numbers!) to return the highest value. For example, if $a=3$, $b=7$ and $c=5$, `challenge_1(a, b, c)` should return the value for $b$, being $7$.

You write your code on the line `# YOUR CODE HERE` which you can delete afterwards. The `NotImplementedError()` makes sure that an empty function does not raise errors. You can also delete this line. 

### Example input
```python
3
7
5
```

### Example return value
```python
7
```

<!-- grade_id: code-chal-01, points: 0, grade: false, solution: true -->
```python
# Code cell: complete the function

def challenge_1(a, b, c):
    # YOUR CODE HERE 
    raise NotImplementedError()
```

If you think your code is correct, you can test your function by running the cell below on a set of numbers. You can start with the numbers above, but you could also try other numbers.

```python
# Test cell

test_challenge_1(challenge_1)
```

Last, the cell below tests your function on a set of numbers (other than the ones you may have used). If you written generic code, you should get the message: `### ADD HERE ###`.

<!-- grade_id: assert-chal-01, points: 10, grade: true, solution: false -->
```python
# Assertion cell

check_assertion(challenge_1)
```

## Next challenge

Great job! If you solved this challenge, you are ready for **Challenge 2a**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge2a.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 2a"/></a>

If you were not able to complete the challenge above, you can practice your skills a bit more by **Challenge 2b**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge2b.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 2b"/></a>