### Before you start...

Before you start solving the challenge, press `Ctrl-Enter` (PC) or `Command-Enter` (Mac) in the cell below. It will load the modules needed for testing and auto-grading the challenges. After you did, a number between the square brackets will appear in front of the cell, like - for example:
```python
[1]
```
If successful, the following message will appear below the first cell:
```python
Assertions imported OK
Test cells imported OK
```

```python
# Run this cell first!

import httpimport
with httpimport.github_repo('eur-nl', 'assertions'):
    from assertions import *
with httpimport.github_repo('jjengelberts', 'precourse-test'):
    from test_cells import *
```

<br>
<br>
<br>
<br>

## Challenge 3b: Traveling In The Netherlands

In this challenge, you will create a function that helps finding the fastest car route from Rotterdam to Apeldoorn, given timings between the relevant highway junctions. A tool like Google Maps will suggest the fastest route, taking current traffic into account. 

A schematic overview of the route is shown in Figure 1. 
<div>
    <img src="https://raw.githubusercontent.com/jjengelberts/precourse-test/main/figures/rotterdam_apeldoorn.svg"><br>
    <b>Figure 1</b> - Schematic overview of highway route Rotterdam-Apeldoorn.
</div>

The numbers of the route represent the following parts (highway numbers in brackets):

* <span style="color:blue">0</span>: Rotterdam-Utrecht (A20 - A12)
* <span style="color:blue">1</span>: Utrecht-Hoevelaken (A27 - A1)
* <span style="color:blue">2</span>: Utrecht-Hoevelaken (A28)
* <span style="color:blue">3</span>: Utrecht-Apeldoorn (A12 - A50)
* <span style="color:blue">4</span>: Hoevelaken-Apeldoorn (A1)

The function `challenge_3b(timings)` has a single argument `timings`, which is a 5 element list with the timings, in minutes, for each numbered road part in Figure 1. So, the time it takes to get from Rotterdam to Utrecht (part <span style="color:blue">0</span>) can be accessed from `timings[0]`, and so on. Finalize `challenge_3b(timings)` in such a way that it returns the fastest timing (in minutes) for the given `timings` list.

### Hints
* From the figure, express all three possible routes in terms of their road part numbers
* Calculate the timing for all of these routes
* Return the lowest of these three numbers

### Example input
```python
[49, 28, 21, 51, 32]
```

### Example return value
```python
100
```
**TIP** Elements of a list can be added together by `+` when they are numbers

**NB** For the example, the shortest route is 49 minutes from Rotterdam to Utrecht (A20-A12, part <span style="color:blue">0</span>) and then in 51 minutes from Utrecht to Apeldoorn (A12-A50, part <span style="color:blue">3</span>), totalling 100 minutes.

<!-- grade_id: code-chal-03b, points: 0, grade: false, solution: true -->
```python
# Code cell

def challenge_3b(timings):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Test cell

test_challenge_3b(challenge_3b)
```

<!-- grade_id: assert-chal-03b, points: 10, grade: true, solution: false -->
```python
# Assertion cell

check_assertion(challenge_3b)
```

## Final remarks

You have reached the end of this Precourse test!

### Score 1 out of 3
If you did not manage to solve this challenge, you have managed to solve one of the least difficult challenges. I would be a good idea to practice a bit more with Python before the course starts.

A good starting point would be the Python course on [Kaggle](https://www.kaggle.com/learn/python). It briefly goes through the Python basics, which might help you to improve your current Python skills.

### Score 2 out of 3
If you did manage to solve this challenge, you have managed to solve two of the challenges. Of course, this is no guarantee since the result depends on the number of Google searches, ChatGPT requests and answers from friends you have used. If you are confident that you did it pretty much on your own, you are prepared for the course.