# A Bit Of Data Science

Before you start solving the challenge, press `Ctrl-Enter` (PC) or `Command-Enter` (Mac) in the cell below. It will load the modules needed for testing and auto-grading the challenges. After you did, a number between the square brackets will appear in front of the cell, like - for example:
```python
[1]
```
If successful, the following message will appear below the first cell:
```python
Assertions imported OK
Test cells imported OK
```

```python
# Run this cell first!

import httpimport
with httpimport.github_repo('eur-nl', 'assertions'):
    from assertions import *
with httpimport.github_repo('jjengelberts', 'precourse-test'):
    from test_cells import *
```

<br>
<br>
<br>
<br>

## Challenge 2b

As a Data Scientist, it is quite common to encounter "dirty" data. For example, you might have a table with a column that should contain numbers, except some cells contain something else.

In this challenge your function `challenge_2b(raw_data)` takes a list corresponding to a column of data where some of the cells contain the word `Missing` (see the example).

Finalize `challenge_3b(raw_data)` in such a way that it calculates the average of the numbers present in the list, while excluding the `Missing` values.

## Example input
```python
[1.0, 2.0, 'Missing', 4.0, 5.0, 'Missing', 'Missing', 8.0, 9.0, 10.0]
```

## Example return value
```python
5.0
```

<!-- grade_id: code-chal-01, points: 0, grade: false, solution: true -->
```python
# Code cell

def challenge_2b(raw_data):
    # YOUR CODE HERE
    raise NotImplementedError()
    return average
```

```python
# Test cell

test_challenge_2b(challenge_2b)
```

<!-- grade_id: assert-chal-01, points: 10, grade: true, solution: false -->
```python
# Assertion cell

check_assertion(challenge_2b)
```

## Next challenge

If you managed to solve this challenge, please do continue to **Challenge 3b**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge3b.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 3b"/></a>

If you were not able to complete the challenge above, please do continue to **Challenge 3c**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge3c.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 3c"/></a>