# Let's Play Chess!

Before you start solving the challenge, press `Ctrl-Enter` (PC) or `Command-Enter` (Mac) in the cell below. It will load the modules needed for testing and auto-grading the challenges. After you did, a number between the square brackets will appear in front of the cell, like - for example:
```python
[1]
```
If successful, the following message will appear below the first cell:
```python
Assertions imported OK
Test cells imported OK
```

```python
# Run this cell first!

import httpimport
with httpimport.github_repo('eur-nl', 'assertions'):
    from assertions import *
with httpimport.github_repo('jjengelberts', 'precourse-test'):
    from test_cells import *
```

<br>
<br>
<br>
<br>

## Challenge 2a
One of the stronger pieces in chess is the rook. Rooks move along rows or columns. In Figure 1, you will see a white rook on D5 with crosses horizontally and vertically indicating the possible squares it can move to. 
<div>
    <img src="https://raw.githubusercontent.com/jjengelberts/precourse-test/main/figures/rookmoves.svg" align="left" width="450"/><br clear="all">
    <b>Figure 1</b>
</div>
<br>

Finalize function `challenge_2a(square)` in such a way that it returns a list `possible` that will contain all possible squares the rook can move to (on an otherwise empty board) given that the rook is presently on the `square` given by the argument. Note that the list is returned in **sorted** order, you don't have to sort it yourself.

**NB**: Looping through rows and columns will also result in adding the current square the rook is on, if you do not take precautions. Be aware of this and make sure **not** to include the given square in your `possible` moves list!


## Example input
```python
D5
```

## Example return value
```python
['A5', 'B5', 'C5', 'D1', 'D2', 'D3', 'D4', 'D6', 'D7', 'D8', 'E5', 'F5', 'G5', 'H5']
```

<!-- grade_id: code-chal-02a, points: 0, grade: false, solution: true -->
```python
# Code cell

def challenge_2a(square):
    col_char = ['A', 'B', 'C', 'D', 'E', 'F', 'G', 'H']
    row_char = ['1', '2', '3', '4', '5', '6', '7', '8']

    # Converting row and column letters to zero-based indices!
    # 'a4' will result in col=0 and row=3
    # 'c6' will result in col=2 and row=5
    col = col_char.index(square[0])
    row = row_char.index(square[1])
    
    possible = []
    # YOUR CODE HERE
    raise NotImplementedError()

    return sorted(possible)
```

```python
# Test cell

test_challenge_2a(challenge_2a)
```

<!-- grade_id: assert-chal-02a, points: 10, grade: true, solution: false -->
```python
# Assertion cell

check_assertion(challenge_2a)
```

## Next challenge

If you managed to solve this challenge, please do continue to **Challenge 3a**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge3a.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 3a"/></a>

If you were not able to complete the challenge above, please do continue to **Challenge 3b**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge3b.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 3b"/></a>