### Before you start...

Before you start solving the challenge, press `Ctrl-Enter` (PC) or `Command-Enter` (Mac) in the cell below. It will load the modules needed for testing and auto-grading the challenges. After you did, a number between the square brackets will appear in front of the cell, like - for example:
```python
[1]
```
If successful, the following message will appear below the first cell:
```python
Assertions imported OK
Test cells imported OK
```

```python
# Run this cell first!

import httpimport
with httpimport.github_repo('eur-nl', 'assertions'):
    from assertions import *
with httpimport.github_repo('jjengelberts', 'precourse-test'):
    from test_cells import *
```

<br>
<br>
<br>
<br>

## Challenge 3c: Reading Literature
John and Mary are big fans of reading books. They both keep track of titles they have read. They compare their lists and are interested in the titles only one of them has read. Finalize the function `challenge_3c(john, mary)` in such a way that it returns a list (**sorted in alphabetical order**) with the titles that either John or Mary has read. The function arguments, `john` and `mary`, are lists with titles read by John and Mary, respectively.

### Example input
**John**
```python
['The Great Gatsby', 'Pride and Prejudice', 'The Catcher in the Rye', 'Wuthering Heights', 'Frankenstein']
```

**Mary**
```python
['Wuthering Heights', 'The Adventures of Huckleberry Finn', 'Jane Eyre', 'Pride and Prejudice', 'Animal Farm']
```

### Example return value
```python
['Animal Farm', 'Frankenstein', 'Jane Eyre', 'The Adventures of Huckleberry Finn', 'The Catcher in the Rye', 'The Great Gatsby']
```

**TIP** To test you can also use simple letters to make it easier on yourself. Consider both the possibility that a book is not read by John as well as a book not read by Mary!


<!-- grade_id: code-chal-03c, points: 0, grade: false, solution: true -->
```python
# Code cell

def challenge_3c(john, mary):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Test cell

test_challenge_3c(challenge_3c)
```

<!-- grade_id: assert-chal-03c, points: 10, grade: true, solution: false -->
```python
# Assertion cell

check_assertion(challenge_3c)
```

## Final remarks

You have reached the end of this Precourse test!

### Score 0 out of 3
If you did not manage to solve this challenge, you have not managed to solve any of the three challenges presented to you. Therefore, we advise you to take a Python course from scratch.

A good starting point is the course on [W3Schools](https://www.w3schools.com/python). This course starts at the beginning and lets you play with the Python basics.


### Score 1 out of 3
If you did manage to solve this challenge, you have managed to solve one of the least difficult challenges. I would be a good idea to practice a bit more with Python before the course starts.

A good starting point would be the Python course on [Kaggle](https://www.kaggle.com/learn/python). It briefly goes through the Python basics, which might help you to improve your current Python skills.