### Before you start...

Before you start solving the challenge, press `Ctrl-Enter` (PC) or `Command-Enter` (Mac) in the cell below. It will load the modules needed for testing and auto-grading the challenges. After you did, a number between the square brackets will appear in front of the cell, like - for example:
```python
[1]
```
If successful, the following message will appear below the first cell:
```python
Assertions imported OK
Test cells imported OK
```

```python
# Run this cell first!
!pip install httpimport
import httpimport
with httpimport.github_repo('eur-nl', 'assertions'):
    from assertions import *
with httpimport.github_repo('jjengelberts', 'precourse-test'):
    from test_cells import *
```

Output:
```
Collecting httpimport
  Downloading httpimport-1.3.1-py2.py3-none-any.whl (17 kB)
Installing collected packages: httpimport
Successfully installed httpimport-1.3.1
Assertions imported OK
```

<br>
<br>
<br>
<br>

## Challenge 3a: Conway's Game Of Life

In this challenge, we are going to use 3x3 list of lists for our version of [Conway's Game of Life](https://en.wikipedia.org/wiki/Conway%27s_Game_of_Life). The Game of Life is a 0-player board game. Each square on the board represents a cell, which can be living or dead. In Figure 1, a 6x6 board is shown with 7 living cells (dark) and 29 dead cells (light).

<table width="100%">
    <tr>
        <td style="text-align: center"><img src="https://raw.githubusercontent.com/jjengelberts/precourse-test/main/figures/98px-Game_of_life_loaf.svg.png"><br><b>Figure 1</b> - showing 7 living and 29 dead cells</td>
        <td style="text-align: center"><img src="https://raw.githubusercontent.com/jjengelberts/precourse-test/main/figures/Game_of_life_animated_glider.gif"><br><b>Figure 2</b> - the glider</td>
    </tr>
</table>

As a player, you set up the board, *i.e.* the initial coloring of the cells. Each "turn" in the game will update the state of each cell. If the number of surroundings living cells is high enough, the cell will survive. To determine whether the cells survive or die, the following rules are applied:

### Rules
1. If a currently living cell has either two or three currently living neighbours, it will remain alive and otherwise it will die.
2. If a currently dead cell has precisely three currently living neighbours, it will be revived to become a living cell and otherwise it will remain dead.

These apparently simple rules actually give rise to some surprisingly complex behaviour. As an example, the so-called _glider_ (Figure 2) is an example of a kind of "organism" (connected bunch of cells) which will move forever, as the rules are applied over and over again, presuming the board were infinite in each direction.

You will look whether the center cell should be alive or dead depending on the neighbouring cells. The 3x3 board is a list of lists with zeros (representing dead cells) and ones (representing living cells). To complete the function `challenge_3a(board)`, you only have to apply the first rule above and return the value `0` if the center cell will be dead in the next generation or `1` if it will be alive. You can follow the steps outlined below:

### Steps
1. Determine whether the center cell (the value at `board[1][1]`) is currently alive or dead
2. Count the number of living cells surrounding the center cell
3. Determine whether the center cell will be alive (`1`) or dead (`0`) in the next generation, based on the rules described above

**TIP**: Use for loops to address all cells on the board. Skip the 1 in the center when you count the number of living cells. 

### Example visual 1
<div>
    <img src="https://raw.githubusercontent.com/jjengelberts/precourse-test/main/figures/dead_alive.png" width="10%"><br>
    <b> Figure 3</b> - Dead cell surrounded by three live cells (will come to life)
</div>

### Example input 1
```python
[[1, 0, 0], [1, 0, 0], [0, 1, 0]]
```  
     
### Example return value 1
```python
1
```

### Example visual 2
<div>
    <img src="https://raw.githubusercontent.com/jjengelberts/precourse-test/main/figures/alive_dead.png" width="10%"><br>
    <b> Figure 4</b> - Living cell surrounded by four living cells (will die)
</div>

### Example input 2
```python
[[1, 0, 1], [0, 1, 0], [1, 0, 1]]
```  
     
### Example return value 2
```python
0
```

<!-- grade_id: code-chal-03a, points: 0, grade: false, solution: true -->
```python
# Code cell

def challenge_3a(board):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Test cell

test_challenge_3a(challenge_3a)
```

Output:
```
Please enter a series of 3 zeroes and ones for row 1, separated by spaces (Press Enter for default values 1 0 0): 1 0 0
Please enter a series of 3 zeroes and ones for row 2, separated by spaces (Press Enter for default values 1 0 0): 1 0 0
Please enter a series of 3 zeroes and ones for row 3, separated by spaces (Press Enter for default values 0 1 0): 0 1 0
0
1
2

Your function returns '1' and hence indicates that the center cell will be alive.
```

<!-- grade_id: assert-chal-03a, points: 10, grade: true, solution: false -->
```python
# Assertion cell

check_assertion(challenge_3a)
```

Output:
```
Call challenge_3a([[1, 0, 0], [1, 0, 0], [0, 1, 0]]) ... OK
Call challenge_3a([[1, 0, 1], [0, 1, 0], [1, 0, 1]]) ... OK
Call challenge_3a([[0, 0, 0], [0, 1, 0], [0, 0, 0]]) ... OK
Call challenge_3a([[1, 1, 1], [1, 0, 1], [1, 1, 1]]) ... OK
Call challenge_3a([[0, 0, 0], [1, 1, 1], [0, 1, 0]]) ... OK
Call challenge_3a([[0, 0, 1], [0, 0, 1], [0, 0, 1]]) ... OK

Well done! You seem to have solved challenge_3a!
```

## Final remarks

You have reached the end of this Precourse test!

### Score of either 2 or 3 out of 3
Depending on whether you solved this challenge, you will have either scored 2 or 3 out of 3 points. Of course, this is no guarantee since the result depends on the number of Google searches, ChatGPT requests and answers from friends you have used. If you are confident that you did it pretty much on your own, you are prepared for the course.