# Off We Go!

You are ready to start the first challenge. The first cell loads the modules needed for testing and auto-grading the challenges. To run the cell, press `Ctrl-Enter` (PC) or `Command-Enter` (Mac) in the cell below. 

If everything is corrrect, a number between the square brackets will appear in front of the cell, like - for example:
```python
[1]
```
If the modules are successfully loaded, the following message will appear below the first cell:
```python
Assertions imported OK
Test cells imported OK
```

```python
# Run this cell to load the modules!

import httpimport
with httpimport.github_repo('eur-nl', 'assertions'):
    from assertions import *
with httpimport.github_repo('jjengelberts', 'precourse-test'):
    from test_cells import *
```

<br>
<br>
<br>
<br>

## Challenge 1: Find highest number

In this first challenge, you will complete a function that finds the highest number. 

The function is called `challenge_1(a, b, c)` and should compare the three arguments `a`, `b` and `c` (these should all be numbers!) to return the highest value. For example, if $a=3$, $b=7$ and $c=5$, `challenge_1(a, b, c)` should return the value for $b$, being $7$.

You should start below the comment line `# YOUR CODE HERE` by replacing the line `raise NotImplementedError()` with your code. The `raise` keyword ensures that an unsolved challenge will raise an error.

### Example input
```python
3
7
5
```

### Example return value
```python
7
```

**TIP** To check what happens inside your function, use `print` messages. Once you run the challenge, you can tell from the print messages whether you do what you intend to do.

<!-- grade_id: code-chal-01, points: 0, grade: false, solution: true -->
```python
# Code cell: complete the function

def challenge_1(a, b, c):
    # YOUR CODE HERE
    raise NotImplementedError()
```


If you think your code is correct, you can test your function by running the cell below on a set of numbers. You can start with the numbers above, but you could also try other numbers.

```python
# Test cell

test_challenge_1(challenge_1)
```

Last, the cell below asserts your function with a set of test cases, other than the ones you may have used. If you have written generic code, you should get an `OK` message on each of them, and after that:
```
Well done! You seem to have solved challenge_1!
```

<!-- grade_id: assert-chal-01, points: 10, grade: true, solution: false -->
```python
# Assertion cell

check_assertion(challenge_1)
```

## Next challenge

Great job! If you solved this challenge, you are ready for **Challenge 2a**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge2a.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 2a"/></a>

If you were not able to complete the challenge above, you can practice your skills a bit more by **Challenge 2b**:

<a href="https://colab.research.google.com/github/jjengelberts/precourse-test/blob/main/Challenge2b.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Challenge 2b"/></a>